---
title: 'String Manipulation'
description: ""
---

## STRING SPLIT

```yaml
type: NormalExercise
key: 3e1f639965
xp: 100
```

The query STRING_SPLIT() separates text with a delimiter

`@instructions`
Take this quote and insert into the STRING_SPLIT() function with this delimiter " ".
"Mysterio's visual on-screen adaptation in #SpiderManFarFromHome is very close to its comic counterpart!"

`@hint`


`@pre_exercise_code`
```{r}
SELECT VALUE FROM  STRING_SPLIT("Mysterio's visual on-screen adaptation in #SpiderManFarFromHome is very close to its comic counterpart!", " ");
```

`@sample_code`
```{r}

```

`@solution`
```{r}
SELECT VALUE FROM  STRING_SPLIT("Mysterio's visual on-screen adaptation in #SpiderManFarFromHome is very close to its comic counterpart!", " ");
```

`@sct`
```{r}

```

---

## REVERSE

```yaml
type: NormalExercise
key: ddeb40f275
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## CONCAT

```yaml
type: NormalExercise
key: 2f125fd068
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## TRIM

```yaml
type: NormalExercise
key: 3bd9f5e0c6
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
