---
title: 'Regular Expressions in SQL Server'
description: ""
---

## IS MATCHING

```yaml
type: NormalExercise
key: 7bba433588
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## SUBSTITION

```yaml
type: NormalExercise
key: 135d522eab
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## WHITESPACE

```yaml
type: NormalExercise
key: c6611a7b8d
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
