---
title: 'Natural Language Processing'
description: ""
---

## TOKENIZATION

```yaml
type: NormalExercise
key: 170446bd58
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## TERM FREQUENCY

```yaml
type: NormalExercise
key: b6532bf985
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## INVERSE DOCUMENT FREQUENCY

```yaml
type: NormalExercise
key: f6134772ed
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## TF-IDF Weighting

```yaml
type: NormalExercise
key: bbc3572230
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
